# STATS 507, Week 1 practice

### Problem 1
Write a Python function that prints "Hello, World!".

```python
def print_hello_world():
    return "Hello World"
    raise NotImplementedError()

print_hello_world()
```

Output:
```
'Hello World'
```

### Problem 2
Write a function that takes two integers as argumens. Return the first argument if it is a positive number. Otherwise, return the sum of both arguments.

<!-- grade_id: cell-db7233a2d95c6406, points: 0, grade: false, solution: true -->
```python
def postive_or_sum(x, y):
    if x > 0:
        return x
    else:
        return x + y
    raise NotImplementedError()

assert postive_or_sum(1,2) == 1
postive_or_sum(1, 2)
postive_or_sum(3, 4)
```

Output:
```
3
```

<!-- grade_id: cell-d525ecbed5402db1, points: 1, grade: true, solution: false -->
```python
assert postive_or_sum(42, 16) == 42
assert postive_or_sum(33, 13) == 33
assert postive_or_sum(-33, 13) == -20
```

```python
assert postive_or_sum(3,4) == 7
```

Output:
```
AssertionError                            Traceback (most recent call last)
Cell In[16], line 1
----> 1 assert postive_or_sum(3,4) == 7

AssertionError:
```

```python

```

```python

```

```python

```