# STATS 507, Week 2_1 practice

```python
# Python primitive
x = None
type(x)

float("1.0")
```

Output:
```
1.0
```

### Problem 1
Write a function that takes two integers as argumens. Return the first argument if it is a positive number. Otherwise, return the sum of both arguments.

<!-- grade_id: cell-db7233a2d95c6406, points: 0, grade: false, solution: true -->
```python
def postive_or_sum(x, y):
    # YOUR CODE HERE
    if x > 0:
        return x
    else:
        return x + y


assert postive_or_sum(1,2) == 1
```

<!-- grade_id: cell-d525ecbed5402db1, points: 1, grade: true, solution: false -->
```python
assert postive_or_sum(42, 16) == 42
assert postive_or_sum(33, 13) == 33
assert postive_or_sum(-33, 13) == -20
```

```python
assert postive_or_sum(3,4) == 3
```

```python

```